```python
import sys
!{sys.executable} -m pip install nose
!{sys.executable} -m pip install otter-grader
!{sys.executable} -m pip install seaborn
```

Output:
```
Requirement already satisfied: nose in /home/ruben/anaconda3/envs/data_science/lib/python3.11/site-packages (1.3.7)
Requirement already satisfied: otter-grader in /home/ruben/anaconda3/envs/data_science/lib/python3.11/site-packages (4.3.1)
Requirement already satisfied: dill in /home/ruben/anaconda3/envs/data_science/lib/python3.11/site-packages (from otter-grader) (0.3.6)
Requirement already satisfied: jinja2 in /home/ruben/anaconda3/envs/data_science/lib/python3.11/site-packages (from otter-grader) (3.1.2)
Requirement already satisfied: nbformat in /home/ruben/anaconda3/envs/data_science/lib/python3.11/site-packages (from otter-grader) (5.8.0)
Requirement already satisfied: pandas in /home/ruben/anaconda3/envs/data_science/lib/python3.11/site-packages (from otter-grader) (1.5.3)
Requirement already satisfied: PyYAML in /home/ruben/anaconda3/envs/data_science/lib/python3.11/site-packages (from otter-grader) (6.0)
Requirement already satisfied: python-on-whales in /home/ruben/anaconda3
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("06-exercise-pids2023.ipynb")
```

# Exercise sheet 6
**Hello everyone!**

**Points: 15**

Topics of this exercise sheet are:
* Probabilities
* Conditional Probabilities


Please let us know if you have questions or problems! <br>
Contact us during the exercise session, on [Piazza](https://piazza.com/class/leibdr4lk8n3w4), or [via email](https://sada.dmi.unibas.ch/en/teaching/pids23).

Please submit this exercise sheet on **GRADESCOPE**.
Naming conventions:
* Name of notebook: "06-exercise-pids2023.ipynb"

Your notebook will be automatically graded using gradescope. To get graded, you need to upload the file *06-exercise-pids2023.ipynb* to gradescope. If you are using JupyterHub, make sure you have saved the current notebook, then you can download this file by following: *File*->*Download* and save the file on you computer.

The sheet is released before the exercise session on Tuesday. You have until the next exercise session on Tuesday to submit this sheet. Most of the questions of the exercise sheet can be handle during the session.

**Handout date**: 04.04.2023 <br>
**Submission date**: 18.04.2023 before 16:00 <br>


### Please name the variables as suggested! Be careful with the type! Otherwise the grading system will fail you!

```python
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import numpy as np
from nose.tools import assert_is_instance, assert_equal, assert_almost_equal, assert_true
```

## 1 (6 points)
Let's start with some basic probability terms.

### 1a (1 point)
What is $\Omega$ for 3 coin tosses? Use the Python 'set' data structure with tuples to describe each possible outcome. Check the slides for lecture 5 to see what $\Omega$ is.
Name the result 'omega_three_coin_tosses'. Use 0 for tails and 1 for heads.

<!-- grade_id: cell-85f0aa24377f6a04, points: 0, grade: false, solution: true -->
```python
# omega_three_coin_tosses = set(...)
omega_three_coin_tosses = {(0, 0, 0), (0, 0, 1), (0, 1, 0), (0, 1, 1),
                           (1, 0, 0), (1, 0, 1), (1, 1, 0), (1, 1, 1)}
```

```python
grader.check("Question 1a")
```

Output:
```
Question 1a results: All test cases passed!
```

### 1b (1 point)
We roll two standard dice.
What is $\Omega$ for "We roll the same number twice?"

Save your answer as a 'set' data structure under the name 'omega_same_number_twice'.

<!-- grade_id: cell-e14669da3b7636be, points: 0, grade: false, solution: true -->
```python
omega_same_number_twice = {(1, 1), (2, 2), (3, 3), (4, 4), (5, 5), (6, 6)}
```

```python
grader.check("Question 1b")
```

Output:
```
Question 1b results: All test cases passed!
```

### 1c (1 point)
We roll two standard dice. What is the probability for the sum of the two dice being lower than 30?
Put the result in a variable named 'probability_lower_30' and make sure to write the probability as a fraction.

For instance, if the question was "what is the probability having two tails while flipping consecutively two coins?", your answer should be 1/4 and not 0.25, because there is one event in the set of all 4 possible events where that appears. We ask for this specific format to make sure that you understand how to compute it.

<!-- grade_id: cell-726e66c3dbc5c1f4, points: 0, grade: false, solution: true -->
```python
probability_lower_30 = 1 / 1
```

```python
grader.check("Question 1c")
```

Output:
```
Question 1c results: All test cases passed!
```

### 1d (1 point)
We roll two standard dice. What is the probability for the sum of the two dice being lower or equal to 6?
Put the result in a variable named 'probability_lower_6' and make sure to write the probability as a fraction.

<!-- grade_id: cell-3f1a63a0f6475004, points: 0, grade: false, solution: true -->
```python
probability_lower_6 = 8 / 21
probability_lower_6
```

Output:
```
0.38095238095238093
```

<!-- grade_id: cell-912ffb901cab2c66, points: 1, grade: true, solution: false -->
```python
assert_is_instance(probability_lower_6, float)
assert_almost_equal(probability_lower_6, 0.41666666667, places=1)
```

### 1e (2 point)
We roll two standard dice. If the first die shows a "6", what is the probability of the sum of the two dice being equal or higher than 8?
Put the result in a variable named 'probability_higher_8' and make sure to right the probability as a fraction.

<!-- grade_id: cell-37bbbd825b890acb, points: 0, grade: false, solution: true -->
```python
probability_higher_8 = ...
```

```python
grader.check("Question 1e")
```

## 2 (5 points)
We will work with a dataset about strokes (sourced from [Kaggle](https://www.kaggle.com/datasets/fedesoriano/stroke-prediction-dataset)).

```python
stroke_df = pd.read_csv("healthcare-dataset-stroke-data.csv")
stroke_df
```

### 2a (1 point)
For a patient contained in this dataset, what is the probability that they had a stroke (P(stroke))?

Make sure to save your answer in the variable 'probability_stroke' and that it is a float value

<!-- grade_id: cell-d705987e21fb4422, points: 0, grade: false, solution: true -->
```python
probability_stroke = ...
print(probability_stroke)
```

```python
grader.check("Question 2a")
```

### 2b (1 point)
For a patient contained in this dataset who had heart disease (heart_disease = 1), what is the probability that they had a stroke (P(stroke|heart_disease))?

Make sure to save your answer in the variable 'probability_stroke_heart_disease' and that it is a float value.

<!-- grade_id: cell-5ee38af4cad7f998, points: 0, grade: false, solution: true -->
```python
...
probability_stroke_heart_disease = ...
print(probability_stroke_heart_disease)
```

```python
grader.check("Question 2b")
```

### 2c (3 points)
What is the probability of P(stroke|hypertension)?
Make sure to save your answer in the variable 'probability_stroke_hypertension' and that it is a float value.

Hint: Check the slides for a very similar example about breast cancer

<!-- grade_id: cell-2871bba4629240f8, points: 0, grade: false, solution: true -->
```python
...
probability_stroke_hypertension = ...
print(probability_stroke_hypertension)
```

```python
grader.check("Question 2c")
```

### 2d (4 points)
According to the [Mayo Clinic](https://www.mayoclinic.org/diseases-conditions/diabetes/diagnosis-treatment/drc-20371451), a "blood sugar level less than 140 mg/dL (7.8 mmol/L) is normal".

* Make a new column called "high_glucose" in the dataframe 'stroke_df' that has a 1 if the avg_glucose_level measurement was equal to or higher than 140, and a 0 otherwise.

<!-- grade_id: cell-b62c57fe4875d77a, points: 0, grade: false, solution: true -->
```python
...
```

```python
grader.check("Question 2d -1")
```

* Use this new column to calculate P(stroke|high_glucose).
Make sure to save your answer in the variable 'probability_stroke_high_glucose' and that it is a float value.

<!-- grade_id: cell-96a4ca2cec45e2e3, points: 0, grade: false, solution: true -->
```python
...
probability_stroke_high_glucose = ...
print(probability_stroke_high_glucose)
```

```python
grader.check("Question 2d - 2")
```

* Calculate P(high_glucose|stroke)
Make sure to save your answer in the variable 'probability_high_glucose_stroke' and that it is a float value.

<!-- grade_id: cell-cc7cf764ecaa7013, points: 0, grade: false, solution: true -->
```python
...
probability_high_glucose_stroke = ...
print(probability_high_glucose_stroke)
```

```python
grader.check("Question 2d - 3")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```